---
title: 'Defining Error, Uncertainty, and Risk'
description: 'This chapter defines error, introduces measures of uncertainty, and provides a framework for evaluating risk.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Defining Error, Uncertainty, and Risk

```yaml
type: VideoExercise
key: 0b52be132e
xp: 50
```

`@projector_key`
83a1ced488978ee6fb5396bd1f5ebc04
